```python
# Initialize Otter
import otter
grader = otter.Notebook("ps1.ipynb")
```

## Docker setup (Instructions applicable to all problem sets)

If you are using Windows operating system, you are encouraged to install Docker and download the necessary image instead of directly downloading the necessary software on your operating system as it is much easier to work with Docker container. 

To use docker follow the instructions below
* Visit https://docs.docker.com/desktop/ and click on Install Docker Desktop that is meant for your Operating System and follow through the instructions.
* Open the Console (a.k.a. Terminal) application on your computer
* Visit https://hub.docker.com/repository/docker/jravi123/umich-507/general and follow the instructions to pull jravi123/umich-507:latest image by running the commands from the Console.

**Note:** If you execute the 'run' command on your image more than once, you will spawn multiple Jupyter Lab servers and there by will notice the second time around it fails to start Jupyter Lab due to port conflict (port 8888 is already occupied by the first server). You can manage these instances by opening the 'Container/Apps' tab of your Docker desktop Dashboard.

# STATS 507
## Problem Set 1

This problem set is designed to introduce to Jupyter notebooks and our autograder platform. All the problems are one-liner answers. Although you may use multiple statements to solve them - it is not required.
 
The first cell imports the autograder package that we are using to autograde your assignments. If running the above cell generates an error, it means that you do not have this package installed. Students who are using Docker will not have this issue as it is already installed. 
Others can install it by uncommenting and running the following command:

```python
# %pip install otter-grader
```

You only need to install this package once. Restart your kernel after installing and verify that the `otter` package loads. 

## Autograded questions
Questions 1-4 are autograded. You should fill in the answers where indicated, and then submit this notebook to Gradescope for grading as detailed in the assignment instructions.

### Question 1
Use the `print` funciton to print "Hello world!!" to the console.

```python
print("Hello world!!")
```

### Question 2
Assign the value of the smallest prime number to a variable 'x'

```python
x = ...
```

```python
grader.check("q2")
```

### Question 3
Write a function `get_mean(x1, x2)` that takes two numerical values `x1` and `x2` and returns the mean of those two numbers. You can assume that the function will be called with only numerical values. 

```python
def get_mean(x1, x2):
    ...
```

```python
grader.check("q3")
```

### Question 4
Assign to the variable `nmin_2024` the number of minutes in the year 2024. This question has one visible and one hidden test.

```python
# Note: 2024 is a leap year
nmin_2024 = ...
```

```python
grader.check("q4")
```

### Question 5

Write a function to convert any given value to float. Assume that your function will be called with only those values that can be converted to 'float'. This problem has 2 visible and 2 hidden tests

```python
def get_float(x):
    ...
```

```python
grader.check("q5")
```

## Manually graded
Questions 6 is manually graded. For PS2 onwards, the graders would look at these responses and grade them by hand. For this problem set, we will use them to collect some basic information about your computing background and interests.

<!-- BEGIN QUESTION -->

### Question 6
How much programming experience do you have?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2 (points: 1)
```
>>> isinstance(x, int)
True
```
```
>>> x == 2
True
```


Tests for q3 (points: 1)
```
>>> assert get_mean(0, 2) == 1
>>> assert get_mean(2, 2) == 2
>>> assert get_mean(4, 8) == 6
```


Tests for q4 (points: 1)
```
>>> isinstance(nmin_2024, int)
True
```


Tests for q5 (points: 1)
```
>>> assert get_float('5.5') == 5.5
>>> assert get_float('-10') == -10
```
