```python
# Initialize Otter
import otter
grader = otter.Notebook("ps8.ipynb")
```

### EDA on crime in Detroit
This data corresponds to the open dataset provided by data.detroitmi.gov
https://data.detroitmi.gov/datasets/detroitmi::rms-crime-incidents/about

In this assignment, you will perform Exploratory Data Analytics on the given dataset. To score full points, you have to provide a conclusion with your insights after exploring the dataset with specific recommendation on what folks could do to avoid being a crime victim in Detroit. Assignments will no conclusion will lose a point.


```python
#import
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import plotly.graph_objects as go
```

### Question 1
Load the dataset from crime_data.csv file given with this file into a dataframe variable named `df`


```python
import pandas as pd
df = pd.read_csv("./crime_data.csv", header=0,
                  dtype={"report_number": str, "arrest_charge": str,
                         "scout_car_area": str, "precinct": str})
```

```python
df.head()
```

```python
grader.check("q1")
```

### Question 2
Get all the column names of the dataframe that has 'offense' mentioned anywhere in the name.
Do not hard code the values. You have to derive the values using a suitable coding expression

```python
cols = ...
...
```

```python
grader.check("q2")
```

### Question 3
Put the name of the only column that has some null values in cols_with_null variable. Here again you have to derive the column name with suitable expression

**Note:** This has only one hidden test.

```python
cols_with_null = ...
...
```

### Question 4
Look into the offense_description column. 
Clean up the offense descriptions by removing the `-` and getting rid of any text after the `-` (along with empty spaces around the word) and then count the number of unique offense descriptions in this dataset.
Assign the number of unique offense description types after cleaning to the variable `count_of_unique_offense_desc`

```python

count_of_unique_offense_desc = ...
...
```

```python
grader.check("q4")
```

### Question 5

Plot the distribution of the 10 most common offense category  and 10 least common offense categories
with appropriate titles for 'x' and 'y' axis in the same figure. To ensure all your tests pass properly, create a figure object  with two axes sub plots `ax`. Assign the `ax[0]` to 10 largest and `ax[1]` to 10 smallest counts



```python
import matplotlib.pyplot as plt

fix, ax = None, None

...
```

```python
grader.check("q5")
```

### Question 6
Write a query to count all the offenses that has the keyword `larceny` in the offense description and assign that value to a variable called `total_larcenies_in_description`. Also get the total number of records that have `larceny` in the offense_category and check if the two counts match. Answer if the counts match or not.

```python

total_larcenies = ...

...
```

### Question 7
Which day of the week has the maximum number of offenses? Write a query to find that and assign the result to `week_day_with_max_offenses`


```python
week_day_with_max_offenses = ...
...
```

```python
grader.check("q7")
```

<!-- BEGIN QUESTION -->

### Question 8
Segment the 24 hours into 
* night from 12 midnight to 5 am
* morning from 5:01 am to 12 noon
* afternoon from 12.01 noon to 5 pm
* evening from 5.01 pm to 12 am

And then show the total crimes for each day of the week and the day segment counts in a suitable chart. 
Hint: You could use one of the seaborn Facet charts for example

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 9
Answer any other questions that might have to come to your mind after working on this dataset for this long. 
You have to show at least 3 different charts


```python
# your answer here
```

```python
# your answer here
```

```python
# your answer here
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 10
Since we have lat/lon information in this dataset, show a suitable geo spatial display of crime locations.

Hint: you could use Plotly Choropleth: https://plotly.com/python/scatter-plots-on-maps/#north-american-precipitation-map


```python
# your answer here
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> assert df is not None
>>> assert isinstance(df, pd.core.frame.DataFrame)
```


Tests for q2 (points: 1)
```
>>> assert all(cols == ['offense_description', 'offense_category', 'state_offense_code'])
```


Tests for q3 (points: 1)


Tests for q4 (points: 1)
```
>>> assert count_of_unique_offense_desc == 98
```


Tests for q5 (points: 2)
```
>>> assert str(ax[0].get_xticklabels()).find('ASSAULT') != -1
>>> assert str(ax[0].get_xticklabels()).find('LARCENY') != -1
>>> assert str(ax[0].get_xticklabels()).find('DAMAGE TO PROPERTY') != -1
>>> assert str(ax[0].get_xticklabels()).find('AGGRAVATED ASSAULT') != -1
>>> assert str(ax[0].get_xticklabels()).find('STOLEN VEHICLE') != -1
>>> assert str(ax[0].get_xticklabels()).find('FRAUD') != -1
>>> assert str(ax[0].get_xticklabels()).find('BURGLARY') != -1
>>> assert str(ax[0].get_xticklabels()).find('WEAPONS OFFENSES') != -1
>>> assert str(ax[0].get_xticklabels()).find('ROBBERY') != -1
>>> assert str(ax[0].get_xticklabels()).find('STOLEN PROPERTY') != -1
```
```
>>> assert str(ax[1].get_xticklabels()).find('GAMBLING') != -1
>>> assert str(ax[1].get_xticklabels()).find('SOLICITATION') != -1
>>> assert str(ax[1].get_xticklabels()).find('SEXUAL ASSAULT') != -1
>>> assert str(ax[1].get_xticklabels()).find('JUSTIFIABLE HOMICIDE') != -1
>>> assert str(ax[1].get_xticklabels()).find('EXTORTION') != -1
>>> assert str(ax[1].get_xticklabels()).find('LIQUOR') != -1
>>> assert str(ax[1].get_xticklabels()).find('KIDNAPPING') != -1
>>> assert str(ax[1].get_xticklabels()).find('FORGERY') != -1
>>> assert str(ax[1].get_xticklabels()).find('OTHER') != -1
>>> assert str(ax[1].get_xticklabels()).find('DISORDERLY CONDUCT') != -1
```


Tests for q6 (points: 2)


Tests for q7 (points: 1)
```
>>> assert week_day_with_max_offenses == 5
```
