```python
# Initialize Otter
import otter
grader = otter.Notebook("ps2.ipynb")
```

# STATS 507
## Problem Set 2
We'll be using the `math` module a few times:

```python
import math
import wrap_test
```

All functions will be tested by visible as well as hidden tests. The maximum amount of time any function is allowed to run is 10 seconds.

### Question 1: Euclid's algorithm 
Euclid's algorithm (<https://en.wikipedia.org/wiki/Euclidean_algorithm>)
is a method for finding the greatest common divisor (GCD) of two
numbers. Recall that the GCD of two numbers $m$ and $n$ is the largest
number that divides both $m$ and $n$.

**1(a)** (4 pts.) The Wikipedia page above includes several pseudocode implementations
    of Euclid's algorithm. Choose one of these, and use it to implement
    a function `gcd(m,n)`, which takes two integers as its arguments and
    returns their GCD. You may assume that both inputs are integers, so
    there is no need to include any error checking in your function. 
    
   **Note**: There are many existing library functions, such as `math.gcd()`, available for calculating this. Do not use them: your job is to implement this function by hand.

```python
def gcd(m, n):
    "Returns the greatest common divisor of integers m and n."
    #solve negative input
    m, n = abs(m), abs(n)
    
    while n != 0:
        m, n = n, m % n
    return m 
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**1(b)** (1 pt.) What does your function do if one or both of its arguments are negative? Does this behavior make sense?

We've identified two approaches for handling potential negative inputs.

1. Since we're interested in finding the GCD, which is always a positive integer, we primarily focus on the magnitude of the numbers. The sign of the inputs doesn't affect the magnitude of the result. Therefore, a straightforward solution is to begin by taking the absolute values of $m$ and $n$ before proceeding with the standard algorithm.

2. Alternatively, we can follow the pseudocode provided in (wiki)[https://en.wikipedia.org/wiki/Euclidean_algorithm], but modify the `return m` statement to `return abs(m)`. To demonstrate its correctness, we can offer the following explanation:

    > In python, the modulo operator % always yields a result with the same sign as its second operand (or zero); the absolute value of the result is strictly smaller than the absolute value of the second operand.

    Note that in `m, n = n, m % n`, the latest $m$ is last $n$ and latest $n$ is remainder of $m/n$, so the latest $m$ and $n$ always share the same sign (both positive or negative). By the calculation method of modulo, we can prove that `m%n == -(-m)%(-n)`. This implies that when both the latest $m$ and $n$ are negative, their updates are unaffected, except for the sign compared to the scenario in which both $m$ and $n$ are positive.  Hence, we only need to change the return value from $m$ to $abs(m)$ to obtain the correct GCD.

<!-- END QUESTION -->

### Question 2: Approximating $e$
The base of the natural logarithm, $e$, is typically defined as the
infinite sum 
$$
e = \sum_{k=0}^\infty \frac{1}{k!}
= 1 + 1 + \frac{1}{2} + \frac{1}{6} + \frac{1}{24} + \dots,$$ where $k!$ denotes the factorial of $k$,
$$
k! = k \cdot (k-1) \cdot (k-2) \cdot \dots \cdot 3 \cdot 2 \cdot 1,
$$
where we define $0!=1$ by convention. For more on Euler's number, see
<https://en.wikipedia.org/wiki/E_(mathematical_constant)>. In this
problem, we will explore different approaches to approximating this
number.


**2(a)** (2 pts.) An early characterization of Euler's number, due to Jacob Bernoulli,
    was as the limit
    $$e = \lim_{x\to\infty} \left(1 + \frac{1}{x}\right)^x.$$ Define a function called `e_limit(n)` that takes as an argument an integer $n$, and returns a `float` that approximates $e$ by taking $x=n$ in the above equation. You may assume that the input to your function will be a positive integer (of course, your function will still run just fine if $n$ is, say, a positive float, but we will only use integer values in what follows).

```python
def e_limit(n):
    "Approximates Euler's number using the limit definition."
    return (1+1/n)**n
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**2(b)** (1 pt.) What happens when you call `e_limit(n)` for really huge values of `n` (say,
    $10^{16}$ or $10^{18}$)? Why does this happen? **Hint:** the answer
    has to do with floating point arithmetic. Think about how $1/n$ is
    represented on your computer when $n$ is big.

For really huge values of `n`, the function is highly likely to return `1.0`. This is due to the way floating-point arithmetic works in Python, following the IEEE 754 standard for representing floating-point numbers.

On a typical 64-bit computer, a floating-point number consists of 1 sign bit, 11 exponent bits, and 53 fraction bits (with one bit implicit). This means that the smallest representable value for $1/n$ (when added to 1) is theoretically $2^{-52}$ or approximately 2.22e-16, as indicated by `sys.float_info.epsilon`.

Therefore, when $1/n$ becomes smaller than this value ($n \sim > 10^{16}$), it results in the loss of precision, causing $1 + 1/n$ to evaluate to 1.0. Consequently, any subsequent calculation involving exponentiation will also yield 1.0.

<!-- END QUESTION -->

**2(c)** (2 pts.) Define a function called `e_series` that takes a single non-negative integer
    argument `n`, and returns an approximation to $e$ based on the first
    `n` terms of the infinite sum shown above. Your function should return a float. You may again assume that the input will be a non-negative integer, so you do not need to include error checking in your function. As an example, `e_series(4)` should return the sum of the first four terms the series, i.e. $1 + 1 + 1/2 + 1/6 \approx 2.667$. 

Notes:

- The sum the equation starts counting with $k=0$ (i.e., it is "$0$-indexed"), while our function starts counting with $n=1$ (i.e., it is "$1$-indexed"). `e_series(1)` should use *one* term, so that `e_series(1) == 1`. Similarly, `e_series(0)` should return $0$, since by convention an empty sum is equal to zero.
- You may use the function `math.factorial(k)` to compute $k!$, but I
    recommend, for the sake of practice, implementing the factorial
    function yourself, as it is a nice example of a problem that is
    easily solved with recursion.

```python
def my_fctl(k):
    return k*my_fctl(k-1) if k != 0 else 1

def e_series(n):
    "Approximates Euler's number using the series definition."
    return sum([1/my_fctl(i) for i in range(n)])
```

```python
grader.check("q2c")
```

Output:
```
q2c results: All test cases passed!
```

**2(d)** (2 pts.) Define a function called `error_table(N)` takes a single positive
    integer `N` as an argument and returns a data structure containing the _relative_ error between the
    approximations given by, respectively, `e_limit` and `e_series`. `error_table(N)` should return a
    list of length $N$. The $k$-th entry of the list, `error_table(N)[k]`, should be a `dict` with two keys, `err_limit` and `err_series`, giving the relative errors for each approach when $n=k+1$:
```
>>> error_table(1)
[{'err_limit': 0.26424111765711533, 'err_series': 0.6321205588285577}]
```
(The first entry of the list is shown.)
**Note**: the true value of $e$, accurate to machine precision, can be obtained as `math.e`.

```python
def error_table(N):
    "Table of relative errors for e_limit(k) and e_series(k) for k=1,...,N."
    return [{
             'err_limit': abs(math.e - e_limit(i+1))/math.e,
             'err_series': abs(math.e - e_series(i+1))/math.e
             } for i in range(N)]
```

```python
grader.check("q2d")
```

Output:
```
q2d results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**2(e)** (2 pt.) 

**Activation Functions**

Later in this course, you will be applying many [activation functions](https://en.wikipedia.org/wiki/Activation_function) to the different layers of your neural network (NN). While the meaning of NN you will learn later, activation functions themselves are simple concepts. In this exercise you will implement LeakyReLU & Tanh activation functions.

**Leaky RELU (rectified linear unit)**

Leaky RELU function is given by:

$$f(x) = {\displaystyle {\begin{cases} 0.01x {\text    {  if }}x\leq 0\\x {\text{  if }}x>0\end{cases}}} $$


**Tanh**

Tanh function a.k.a hyperbolic tangent is given by

$${\displaystyle \tanh(x)\doteq {\frac {e^{x}-e^{-x}}{e^{x}+e^{-x}}}}$$

Although you have derived the value of 'e' in the earlier exercise, you will be using the standard **math** module to compute the value of $e^x$. 

Please complete the leaky_relu and tanh functions below to return values appropriately:

```python
def leaky_relu(x):
    x = float(x)
    return 0.01*x if x<=0 else x
    
```

```python
def tanh(x):
    x = float(x)
    return (1-2/(math.expm1(2*x)+2))
```

```python
grader.check("q2e")
```

Output:
```
q2e results: All test cases passed!
```

<!-- END QUESTION -->

### Question 3: Root finding and minimization
Recall that a *root* $r$ of a function $f:\mathbb{R}\to\mathbb{R}$ satisfies $f(r)=0$. The [bisection method](https://en.wikipedia.org/wiki/Bisection_method) is a generic algorithm for finding the root of a continuous, real-valued function. 

**3(a)** (4 pts.) Write a function `find_root(f, a, b)` which uses the bisection method to locate a root of the function $f$ on the interval $[a,b]$. Here, $f$ is a Python function which accepts a single float, and returns a float (i.e., a real-valued function of one variable.) You may assume that $f$ is continuous and that $f(a)$ and $f(b)$ have different signs, that is, `sgn(f(a)) != sgn(f(b))` where `sgn(x)` is the sign function:

```python
def sgn(x):
    if x < 0:
        return -1
    if x > 0:
        return 1
    return 0
```

(**Note**: due to finite precision issues mentioned above, your algorithm may not find an exact root of $f$, i.e. a number such that `f(r) == 0.0`. We will say that $r$ is a root if $|f(r)|<10^{-16}$.)

```python
def find_root(f, a, b):
    "Find a root of f(x) in the interval [a, b], assuming that sgn(f(a)) != sgn(f(b))."
    if a>b:
        a, b = b, a
    while True:
        mid = (a+b)/2
        if abs(f(mid)) < 1e-16:
            return mid
        if sgn(f(a))*sgn(f(mid)) < 0:
            b = mid
        else:
            a = mid
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**3(b)** (2 pts.) Suppose that $f:[a,b]\to\mathbb{R}$ is a strictly convex and continuously differentiable function. Let $x^*\in [a,b]$ be such that $f(x^*) \le f(x)$ for all $x\in[a,b]$. Explain using calculus how to use the root finding algorithm you implemented above to find $x^*$.

(1) By [extreme value theorem](https://en.wikipedia.org/wiki/Extreme_value_theorem), as $f$ is continuously differentiable, we claim that the $x^*$ described above must exist.     

(2) By (1) and property of strictly convex, we must have one unique $x^*$. (ref: [click](https://math.stackexchange.com/questions/1435502/question-about-existance-of-global-minimum-in-a-strictly-convex-function))     

(3) There are three choices for $x^*$: $x^* \in (a,b)$, $x^*=a$ and $x^*=b$.     

(4) If $x^* \in (a,b)$, then $f'(x^*)=0$. So we can solve $f'(x)=0$ to get $x^*$. However, we need to show that we can get one unique solution of $f'(x)=0$.    

(4.1) By continuously differentiability of $f$, $f'$ is continuous on $[a,b]$. By strictly convexity of $f$, $f'$ is strictly increasing (ref: [click](https://proofwiki.org/wiki/Real_Function_is_Strictly_Convex_iff_Derivative_is_Strictly_Increasing) ). Since $f'$ is continuous and strictly increasing, if the root exists, it must be unique.    

(5) By (4) and (4.1), we solve $f'(x)=0$ on $x\in (a,b)$. If we get the root, then retrun it.    

(6) If we don't get root on $(a,b)$. Then $x^*$ must be either $a$ or $b$. In (4.1) we show that $f'$ is stricly increasing. So,    

(6.1) If $f'(a)<0$, in the neighbor of $a$, we must have $f'<0$, then f decrease in the neighbor, so the minimum point can't be $a$, i.e. $x^*=b$.       
(6.2) If $f'(a)\geq 0$, $f'(x)>f'(a) \geq 0$, f strictly increases, so $x^*=a$.      


<!-- END QUESTION -->

**3(c)** (4 pts.) Write a function `minimize(f_prime, a, b)` that minimizes a function $f$ satisfying the assumptions of the previous question. Here `f_prime` is a function that returns the derivative of $f$ evaluated at the point $x$.

```python
def minimize(f_prime, a, b):
    """
    Minimizes the convex function f over the interval [a,b].
    
    Args:
        f_prime: a function representing the (continuous) derivative of f.
        a, b: interval on which f is defined.
        
    Returns:
        The point x at which f is minimized.
    """
    if sgn(f_prime(a))*sgn(f_prime(b))<0:
        return find_root(f_prime, a, b)
    else:
        return b if f_prime(a)<0 else a
```

```python
grader.check("q3c")
```

Output:
```
q3c results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1a (points: 4)
```
>>> gcd(0, 0) == 0
True
```
```
>>> for m in range(1, 10):
...     assert gcd(0, m) == m
```
```
>>> gcd(1, 1) == 1
True
```
```
>>> for m in range(-10, 10):
...     assert gcd(m, 1) == 1
...     assert gcd(m, 2) == 2 - (m % 2)
```


Tests for q2a (points: 2)
```
>>> e_limit(1) == 2.
True
```
```
>>> e_limit(2) == 2.25
True
```


Tests for q2c (points: 2)
```
>>> e_series(0) == 0
True
```
```
>>> e_series(1) == 1.
True
```


Tests for q2d (points: 2)
```
>>> for k in range(1, 5, 10):
...     assert len(error_table(k)) == k
```
```
>>> assert abs(error_table(1)[0]['err_limit'] - 0.26424111765711533) < 1e-6
>>> assert abs(error_table(1)[0]['err_series'] - 0.6321205588285577) < 1e-6
```


Tests for q2e (points: 2)
```
>>> assert leaky_relu(1) == 1
>>> assert leaky_relu(0) == 0
>>> assert leaky_relu(-1) == -0.01
>>> 
>>> assert tanh(1) == 0.7615941559557649
>>> assert tanh(100) == 1
>>> assert tanh(-100) == -1
```


Tests for q3a (points: 4)
```
>>> import wrap_test
>>> def test4a1():
...     f = lambda x: x ** 3
...     r = find_root(f, -4, 3)
...     assert abs(f(r)) < 1e-16
>>> wrap_test.wrap_test(test4a1)
```
```
>>> import wrap_test
>>> def test4a2():
...     r = find_root(math.cos, 0, 3.14)
...     assert abs(math.cos(r)) < 1e-16
>>> wrap_test.wrap_test(test4a2)
```


Tests for q3c (points: 4)
```
>>> import wrap_test
>>> def test4c1():
...     f_prime = lambda x: 2 * x  # derivative of f(x) = x^2
...     assert abs(minimize(f_prime, -10, 10) - 0) < 1e-8
>>> wrap_test.wrap_test(test4c1)
```
```
>>> import wrap_test
>>> def test4c2():
...     f_prime = lambda x: -math.cos(x)  # derivative of f(x) = -sin(x)
...     assert abs(minimize(f_prime, 0, 3.14) - math.pi/2) < 1e-8
>>> wrap_test.wrap_test(test4c2)
```
