```python
# Initialize Otter
import otter
grader = otter.Notebook("ps5.ipynb")
```

### Question 1
In this question you will implement a class called `Matrix`. A matrix will be modeled as a collection of vectors all of the same dimension, and you may use the `Vector` class that we explored in class (included with this problem set as `vector.py`) to implement it:

```python
from vector import Vector
```

```python
class Matrix:
    '''
    Matrix class
    '''


    def __init__(self, *args) -> None:
        '''
        constructor with args
        '''
        from vector import Vector

        if not args:
            raise ValueError("Matrix should require at least one vector for initialization")
        
        if not all(isinstance(arg, Vector) for arg in args):
            raise TypeError("Input arguments should be vectors")
        
        if not all(args[0].dim==arg.dim for arg in args):
            raise ValueError("Vectors should all have the same dimension")
        self._val = tuple(args)

    def __str__(self) -> str:
        if self.dims[1] == 1:
            return "\n".join('('+str(v)+')' for v in self._val[0])
        return "\n".join(map(str, zip(*self._val)))
    
    def __repr__(self) -> str:
        return str(self)
        
    def __iter__(self):
        return iter(self._val)
    
    def __add__(self, other):
        #if not isinstance(other, Matrix):
        #   return NotImplemented
        
        if not self._conformable(other):
            raise ValueError("Cannot add matrices of different dimensions")
        return Matrix(*[a + b for a, b in zip(self, other)])

    def __mul__(self, other):
        import numbers
        if not isinstance(other, numbers.Number):
            return NotImplemented
        return Matrix(*(a * other for a in self._val))
    
    def __rmul__(self, other):
        return self.__mul__(other)
    
    
    def __radd__(self, other):
        return self.__add__(other)

    def __getitem__(self, items):
        if (len(items)==2)and all(isinstance(i, int) for i in items) and 0<=items[0]<self.dims[0] and 0<=items[1]<self.dims[1]:
                return self._val[items[1]][items[0]]
        else:
            raise IndexError("Illegal Index")
    
    def __matmul__(self, other):
        from vector import Vector
        if not self._conformable_mat(other):
            raise ValueError("Incompatible matrix dimensions")

        newcols = []
        for col_j in other:
            newcol_j = [0.0]*self.dims[0]
            for row_i in range(self.dims[0]):
                newcol_j[row_i] = sum(ik * kj for ik, kj in zip((self[row_i, k] for k in range(self.dims[1])), col_j))
            
            newcols.append(Vector(*newcol_j))

        return Matrix(*newcols)
    
    def __eq__(self, __value: object) -> bool:
        if self._conformable(__value):
            return all((s == v  for s,v in zip(self._val, __value._val)))
        return False

    @property
    def dims(self) -> tuple:
        return (self._val[0].dim, len(self._val))

    def _conformable(self, other) -> bool :
        return isinstance(other, Matrix) and self.dims == other.dims
    
    def _conformable_mat(self, other) -> bool :
        '''conformable 4 matmul'''
        return isinstance(other, Matrix) and self.dims[1] == other.dims[0]
    
    def transpose(self):
        from vector import Vector
        return Matrix(*map(lambda v: Vector(*v), zip(*self._val)))
    
```

**1(a)** (1 pt) The constructor of `Matrix` should accept one or more `Vector` objects, which will represent the *columns* of the matrix.

```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(5, 6, 7)
>>> v3 = Vector(8, 9, 10)
>>> Matrix(v1)
<__main__.Matrix at 0x10efbabb0>
>>> Matrix(v1, v2)
<__main__.Matrix at 0x117023640>
>>> Matrix(v1, v2, v3)
<__main__.Matrix at 0x117023040>
```

Additionally, it should perform the following checks:

1.  It should verify that each argument is of the correct type:
    ```
    >>> Matrix(1, "", v1)
    TypeError: Input arguments should be vectors
    ```
    
2.  It should verify that each argument has the same dimension.

    ```
    >>> v1 = Vector(1, 2, 3)
    >>> v2 = Vector(5, 6)
    >>> Matrix(v1, v2)
    ValueError: Vectors should all have the same dimension
    ```

```python
grader.check("1a")
```

Output:
```
1a results: All test cases passed!
```

**1(b)** (2 pts) Calling `str` on a `Matrix` object should produce a nice looking text representation:

```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> v3 = Vector(7, 8, 9)
>>> M = Matrix(v1, v2, v3)
>>> print(M)
(1, 4, 7)
(2, 5, 8)
(3, 6, 9)
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

**1(c)** (1 pt.) A matrix should have a attribute `dims` which returns a tuple of its dimensions (number of rows and columns):
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> M.dims
(3, 2)
```

```python
grader.check("q1c")
```

Output:
```
q1c results: All test cases passed!
```

**1(d)** (1 pt.) A matrix should be iterable. Iterating over the matrix should return each of its columns as vectors.
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> for col in M: print(col)
(1, 2, 3)
(4, 5, 6)
```

```python
grader.check("q1d")
```

Output:
```
q1d results: All test cases passed!
```

**1(e)** (2 pts) Two matrices of the same dimensions can be added together. 
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> print(M + M)
(2, 8)
(4, 10)
(6, 12)
```
However, matrices of different dimensions cannot be added together:
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> Matrix(v1) + Matrix(v1, v2)
ValueError: Cannot add matrices of different dimensions
```

```python
grader.check("q1e")
```

Output:
```
q1e results: All test cases passed!
```

**1(f)** (2 pts) Matrices can be multipled by a numerical constant:
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> print(M * 2.5)
(2.5, 10.0)
(5.0, 12.5)
(7.5, 15.0)
>>> print("a" * M)
TypeError: ...
```

```python
grader.check("q1f")
```

Output:
```
q1f results: All test cases passed!
```

**1(g)** (2 pts) Individual entries of a matrix can be accessed using the notation `M[i,j]`:
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> print(M)
>>> M[0, 0]
1
>>> M[2, 1]
6
>>> M[3, 3]
IndexError: ...
```

```python
grader.check("q1g")
```

Output:
```
q1g results: All test cases passed!
```

**1(h)** (3 pts) A matrix can be transposed:
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> print(M)
(1, 4)
(2, 5)
(3, 6)
>>> print(M.transpose())
(1, 2, 3)
(4, 5, 6)
```

```python
grader.check("q1h")
```

Output:
```
q1h results: All test cases passed!
```

**1(i)** (2 pts) Two matrices of conformable dimensions can be multiplied using the matrix multiplication operator (`@`). Hint: which dunder method should you implement for this symbol?
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> print(M @ M.transpose())
(17, 22, 27)
(22, 29, 36)
(27, 36, 45)
>>> M @ M
ValueError: Incompatible matrix dimensions
```

```python
grader.check("q1i")
```

Output:
```
q1i results: All test cases passed!
```

**1(j)** (2 pts) Matrices can be tested for equality:
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> M == M
True
>>> M == Matrix(v1)
False
>>> M == 1.0
False
>>> M == "matrix"
False
```

```python
grader.check("q1j")
```

Output:
```
q1j results: All test cases passed!
```

### Question 2
This question will use the MapReduce concepts for data processing.

```python
from functools import reduce
import itertools
import gzip


class MapReduce:
    @property
    def reduce_init(self):
        # override as necessary if the init parameter needs to change
        return None

    def mapper(self, x):
        raise NotImplementedError()

    def reducer(self, accum, x):
        raise NotImplementedError()

    def postprocess(self, reduced):
        # override if necessary
        return reduced

    def run(self, iterable):
        mapped = map(self.mapper, iterable)
        reduced = reduce(self.reducer, mapped, self.reduce_init)
        processed = self.postprocess(reduced)
        return processed
```

(To make things more flexible, we have also added an optional `.postprocess()` method that can be used to do additional processing after the reduction step.)

Questions 2(a) & 2(b) concern the Enron dataset:

```python
def enron(n=None):
    i1 = gzip.open("email-Enron.txt.gz", "rt")
    i2 = itertools.islice(i1, 4, None)  # slice off header
    return itertools.islice(i2, n)
```


For each question below, implement a subclass of MapReduce such that calling `.run(enron(n))` produces the desired output. For example, if the question asked you to calculate the total number of e-mails, your solution could be:

```python
class NumEmails(MapReduce):
    @property
    def reduce_init(self):
        return 0

    def mapper(self, x):
        return 1

    def reducer(self, accum, x):
        return accum + x


NumEmails().run(enron(100))
```

Output:
```
100
```

**2(a)** (3 pts) Define a user's *importance* to be the number of unique people who e-mailed them (not including themself). Write a MapReduce class that returns a `collections.Counter` mapping each user ID to their importance when run.

```python
class Importance(MapReduce):
    '''
    importance = the number of unique people who e-mailed them (not including themself).
    '''
    @property
    def reduce_init(self):
        from collections import defaultdict
        return defaultdict(set)

    def mapper(self, x):
        return tuple(map(int, x.split()))


    def reducer(self, accum, x):
        if x[0] != x[1]:
            accum[x[1]].add(x[0]) 
        return accum

    def postprocess(self, reduced):
        from collections import Counter

        return Counter({k: len(v) for k, v in reduced.items()})
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

**2(b)** (4 pts) Define a user's *forgetfulness* to be the number of times they e-mailed themself. Write a MapReduce class that returns a `Counter` that maps each user who e-mailed themself at least once to their forgetfulness score.

```python
class Forgetful(MapReduce):
    '''Forgetful = the number of times they e-mailed themself'''
    @property
    def reduce_init(self):
        return dict()

    def mapper(self, x):
        nodes = tuple(map(int, x.split()))
        return (nodes[0]==nodes[1], *nodes)

    def reducer(self, accum, x):
        '''x = (flag, send, receiver)'''
        if x[0]:
            accum[x[1]] = accum.get(x[1], 0)+1
        return accum

    def postprocess(self, reduced):
        from collections import Counter

        return Counter(reduced)
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for 1a (points: 1)
```
>>> from vector import Vector
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(5, 6, 7)
>>> v3 = Vector(8, 9, 10)
>>> assert Matrix(v1)
>>> assert Matrix(v1, v2)
>>> assert Matrix(v1, v2, v3)
```
```
>>> try:
...     Matrix()
... except:
...     pass
... else:
...     raise Exception("Matrix should require at least one vector for initialization")
```
```
>>> v1 = Vector(1, 2, 3)
>>> try:
...     Matrix(1, '', v1)
... except TypeError:
...     pass
... else:
...     raise Exception("Should raise a Type Error")
```
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(5, 6)
>>> try:
...     Matrix(v1, v2)
... except ValueError:
...     pass
... else:
...     raise Exception("Should raise a Value Error")
```
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> try:
...     Matrix(v1, v2, [1,2,3])
... except TypeError:
...     pass
... else:
...     raise Exception("Should raise a Type Error")
```


Tests for q1b (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> v3 = Vector(7, 8, 9)
>>> M = Matrix(v1, v2, v3)
>>> assert str(M).strip() == '(1, 4, 7)\n(2, 5, 8)\n(3, 6, 9)'
```


Tests for q1c (points: 1)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert M.dims == (3, 2)
```
```
>>> v1 = Vector(1, 2, 3)
>>> M = Matrix(v1)
>>> assert M.dims == (3, 1)
```


Tests for q1d (points: 1)
```
>>> v = [Vector(1, 2, 3), Vector(4, 5, 6)]
>>> M = Matrix(v[0], v[1])
>>> count = 0
>>> for col in M:
...     assert isinstance(col, Vector)
...     assert col == v[count]
...     count += 1
```


Tests for q1e (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert isinstance(M + M, Matrix)
>>> assert str(M + M).strip() == '(2, 8)\n(4, 10)\n(6, 12)'
```
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> try:
...     Matrix(v1) + Matrix(v1, v2)
... except ValueError:
...     pass
... else:
...     raise Exception('Should raise a value error')
```


Tests for q1f (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert str(M * 2.5).strip() == '(2.5, 10.0)\n(5.0, 12.5)\n(7.5, 15.0)'
>>> assert str(2.5 * M).strip() == '(2.5, 10.0)\n(5.0, 12.5)\n(7.5, 15.0)'
```
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> try:
...     'a' * M
... except TypeError:
...     pass
... else:
...     raise Exception('Should raise a type error')
```


Tests for q1g (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert M[0,0] == 1
>>> assert M[2, 1] == 6
>>> try:
...     M[3, 3]
... except IndexError:
...     pass
... else:
...     raise Exception('Should raise an index error')
```


Tests for q1h (points: 3)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert str(M.transpose()).strip() == '(1, 2, 3)\n(4, 5, 6)'
```


Tests for q1i (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert str(M @ M.transpose()).strip() == '(17, 22, 27)\n(22, 29, 36)\n(27, 36, 45)'
```
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> try:
...     M @ M
... except ValueError:
...     pass
... else:
...     raise Exception('Should raise a value error')
```


Tests for q1j (points: 2)
```
>>> v1 = Vector(1, 2, 3)
>>> v2 = Vector(4, 5, 6)
>>> M = Matrix(v1, v2)
>>> assert M == M
```


Tests for q2a (points: 3)
```
>>> assert isinstance(Importance().run(enron(10)), dict)
```
```
>>> assert Importance().run(enron(10))[0] == 1
```
```
>>> assert Importance().run(enron(10))[1] == 1
```
```
>>> assert Importance().run(enron(1000))[1] == 56
```


Tests for q2b (points: 4)
```
>>> from collections import Counter
>>> 
>>> assert isinstance(Forgetful().run(enron(1000)), Counter)
```
```
>>> assert len(Forgetful().run(enron(1000))) == 1
```
```
>>> assert Forgetful().run(enron(1000))[46] == 3
```
```
>>> assert 97 not in Forgetful().run(enron(10000))
```
